# Homework #2

## PHYS 2600, Spring 2019

__Important notice:__ All cells in your notebook will be run, start to finish, using a fresh kernel when it is graded!  To make sure the graded notebook looks like what you expect, we recommend selecting "Kernel > Restart & Run All" in the menu above before you submit.

If you worked collaboratively on this assignment, __include the names of your collaborators in the cell below:__

<!-- grade_id: cell-6aad98798978a240, points: 0, grade: true, solution: true -->
_Type your answer here using Markdown._

## 2.0 - Homework correction (3 points)

In the cell below, explain and correct __one mistake__ you made on the previous homework assignment.  If you got a perfect score, great!  Tell us what your favorite problem was, or use the space to give any other feedback you might have on the class/tutorials/homework.

<!-- grade_id: cell-1689332322e959be, points: 3, grade: true, solution: true -->
I only got one problem wrong, because I forgot to save my work after I finished it...Hope I have time to do it again, that was a easy one. 

## 2.1 - Fixed-point binary and precision (9 points)

Let's return to the example we saw in lecture of a fixed-point binary representation in 7 bits:

<img src="https://physicscourses.colorado.edu/phys2600/phys2600_fa18/img/fixed-prec.png" />

(Assume we are only dealing with __positive__ numbers here, so there is no sign bit.)

What are the smallest and largest _binary_ numbers that can be represented in this way?  Now what are those numbers in decimal notation?  (The difference between them is the _dynamic range_ of our representation.)

<!-- grade_id: cell-d6069000132514b5, points: 2, grade: true, solution: true -->
smallest is 00000000, largest is 11111111

Let's evaluate the following expression:
\\[
1.1 + 3.2 = ?
\\]

using our binary fixed-point representation, step by step.  First, translate each decimal number to the closest binary fixed-point number.

<!-- grade_id: cell-3aa85982a4332363, points: 2, grade: true, solution: true -->
1.1 = 1*2^0+(2^-1)*0+0*(2^-2)+0*(2^-3)+1*(2^-4)+1*(2^-5)+0*(2^-6)
1.000110
3.2 = 1*(2^1)+1*(2^0)+(2^-1)*0+0*(2^-2)+1*(2^-3)+1*(2^-4)+0*(2^-5)+0*(2^-6)
11.001100

Next, add the two binary numbers from above together _by hand_.

<!-- grade_id: cell-f1f4ed92d3175530, points: 1, grade: true, solution: true -->
100.010011

Finally, translate your result back to decimal from binary fixed-point, and calculate the difference between your answer and the exact result.  (This difference is known as the __round-off error__.)

<!-- grade_id: cell-4d6884f5593b5126, points: 2, grade: true, solution: true -->
4.296875, 
the difference is 0.003125

To help you check your work, note that the largest possible round-off error is exactly 1/8.  This is because the worst-case number for rounding is 1/16: we either have to set it to 0 or 1/8, and any other number in that range would be closer to one or the other.  If, say, we round 1/16 up to 1/8, then 1/16 + 1/16 becomes 1/8 + 1/8 = 1/4, which is 1/8 larger than the right answer.

Instead of just dealing with pairs of numbers at once, suppose we're using a more complicated algorithm that has to iterate many times, like the square root example we saw.  What is the __accumulated round-off error__ if we compute the sum `1.2 + 2.2` one hundred times in a row, using our fixed-point representation?

<!-- grade_id: cell-60938be06a624f46, points: 2, grade: true, solution: true -->
We will have 97 digits to the right of the decimal point after 100 hundred times, since we have 3 digits to the left of decimal. Therefore, the round off error is 2^-97.

## 2.2 - Mistakes and formatting (6 points)

(From exercise 1.9c in Langtangen.)

Consider the following program, which tries to verify that these two equations are correct:

\\[
(a+b)^2 = a^2 + 2ab + b^2 \\
(a-b)^2 = a^2 - 2ab + b^2
\\]

for the values $a=3.3$ and $b=5.3$.

```python
a=3,3   b=5,3
a2 = a**2
b2 = b**2

eq1_sum = a2 + 2ab + b2
eq2_sum = a2 - 2ab - b2

eq1_pow = (a + b)**2
eq2_pow = (a - b)**2

print('First equation: %g = %g', % (eq1_sum, eq1_pow))
print('Second equation: %h = %h', % (eq2_sum, eq2_pow))
```

There are tons of mistakes in this program!  Copy the code into the cell below, then __fix the mistakes__ and run to verify the two equations.

<!-- grade_id: cell-73dc6f433ff4bab7, points: 6, grade: true, solution: true -->
```python
## YOUR CODE HERE
a=3.3   
b=5.3

a2 = a**2
b2 = b**2

eq1_sum = a2 + 2*a*b + b2
eq2_sum = a2 - 2*a*b + b2

eq1_pow = (a + b)**2
eq2_pow = (a - b)**2

print('First equation: %g = %g' % (eq1_sum, eq1_pow))
print('Second equation: %g = %g' % (eq2_sum, eq2_pow))
```

Output:
```
First equation: 73.96 = 73.96
Second equation: 4 = 4
```

## 2.3 - Some exercises with NumPy (10 points)

The following are small problems in working with the `numpy` module and its arrays.

### Part A (2 points)

__Create an array `cos_squared`__, which contains the value of the function $\cos^2(t)$, for $t$ values ranging from 0 to 6 (including both ends!) and spaced by 0.5.  (So `t = [0, 0.5, 1.0, ..., 6.0].`)

<!-- grade_id: cell-67bdf348d2e57ed0, points: 0, grade: false, solution: true -->
```python
import numpy as np
from math import cos, radians
## YOUR CODE HERE
t = np.arange(0,6.2,0.5)
cos = np.cos(t)
cos_squared = np.square(cos)
```

<!-- grade_id: cell-15f2a517158f210f, points: 2, grade: true, solution: false -->
```python
import numpy.testing as npt

print(cos_squared)

assert len(cos_squared) == 13
assert type(cos_squared) is np.ndarray

```

Output:
```
[1.         0.77015115 0.29192658 0.00500375 0.17317819 0.64183109
 0.98008514 0.87695113 0.42724998 0.04443487 0.08046424 0.50221285
 0.92192698]
```

### Part B (2 points)

__Create an array `length_40_array`__ which, as the name implies, should have 40 entries: the first 20 entries should be the number 20, and the second 20 entries should be the number zero.

_(Hint: there are quite a lot of ways to do this!  You are allowed to use a for loop this time, but as we'll see for loops are generally discouraged with numpy.  The `np.concatenate` function might be worth a look...)_

<!-- grade_id: cell-ff71b410dc365eb6, points: 0, grade: false, solution: true -->
```python
import numpy as np

## YOUR CODE HERE
array1 = np.full((1, 20), 20)
array2 = np.full((1, 20), 0)

length_40_array = np.concatenate((array1, array2), axis=None)
```

<!-- grade_id: cell-6ac731fab9cf9e6b, points: 2, grade: true, solution: false -->
```python
import numpy.testing as npt

print(length_40_array)

assert len(length_40_array) == 40
assert type(length_40_array) is np.ndarray

```

Output:
```
[20 20 20 20 20 20 20 20 20 20 20 20 20 20 20 20 20 20 20 20  0  0  0  0
  0  0  0  0  0  0  0  0  0  0  0  0  0  0  0  0]
```

### Part C (3 points)

A _parallelepiped_ is a three-dimensional shape which is a cousin of the cube, except that its three pairs of identical faces are parallelograms - see [the Wikipedia page](https://en.wikipedia.org/wiki/Parallelepiped) for a sketch and in-depth explanation.  If we describe a parallelepiped in terms of the three vectors $\mathbf{a}, \mathbf{b}, \mathbf{c}$ describing its edges at one corner, then its volume can be computed from the _triple product_

\\[
V = |\mathbf{a} \cdot (\mathbf{b} \times \mathbf{c})|.
\\]

__Implement the function `triple_product(a,b,c)` below__ using `numpy`, and then run the test cell to calculate the area of a parallelopiped described by the given vectors.

<!-- grade_id: cell-097330b1f16e6e8d, points: 0, grade: false, solution: true -->
```python
import numpy as np

def triple_product(a=a,b=b,c=c):
    ## YOUR CODE HERE
 
    crossbc = np.cross(b, c)
    return np.dot(a, crossbc)
```

<!-- grade_id: cell-ec8460229817e27f, points: 3, grade: true, solution: false -->
```python
import numpy.testing as npt

# First test: a cube!  Volume should be 1.
a = [1,0,0]
b = [0,1,0]
c = [0,0,1]

print(triple_product(a,b,c))
npt.assert_allclose(triple_product(a,b,c), 1)

# Second test: a rectangular prism!  Volume should be 4.
a = [2,0,0]
b = [0,2,0]
c = [0,0,1]

print(triple_product(a,b,c))
npt.assert_allclose(triple_product(a,b,c), 4)

# Third test: all angles are 45 degrees, a "rhombohedron" with unit side length
# Volume should be 0.455090, to six digits - used formula from Wikipedia.
a = [1,0,0]
b = [0.70710678, 0.70710678, 0.        ]
c = [0.70710678, 0.29289322, 0.64359425]

print(triple_product(a,b,c))
npt.assert_allclose(triple_product(a,b,c), 0.455090, atol=1e-6)



```

Output:
```
1
4
0.45508985774401506
```

### Part D (3 points)

The list `photo_stop_volts` below contains a set of measurements of the stopping voltage $V_{\rm stop}$ (in volts) for a few different materials in a photoelectric effect experiment. (Note that __the stopping voltage is always negative__ in our convention.)  The experiment used ultraviolet light, with a frequency $f$ of 1640 terahertz (THz).  __Create a new list `photo_work_functions`__ that contains the work function $\Phi$ (__in electron-volts, eV__) for each material.

To save you some trouble, here's the relevant equation for the photoelectric effect - you're on your own to look up the physical constants and unit conversions.

\\[
\Phi = hf + eV_{\rm stop} \\
\\]

Here $h$ is Planck's constant (not the reduced constant $\hbar$!) and $e$ is the charge of a single electron.

_(Hint: working in electron-volts instead of SI units will make your life easier...)_

<!-- grade_id: cell-37cb96e4ca9d995e, points: 0, grade: false, solution: true -->
```python
import numpy as np

photo_stop_volts = np.array([-4.50, -1.77, -2.70, -3.88])
## YOUR CODE HERE
h = 4.1357*10**(-15) #eV s
f = 1640*10**12 #Hz

photo_work_functions = h*f +photo_stop_volts

```

<!-- grade_id: cell-2a87e1ce2f886580, points: 3, grade: true, solution: false -->
```python
import numpy.testing as npt

print("Work functions: ", photo_work_functions)

assert len(photo_work_functions) == 4
assert type(photo_work_functions) is np.ndarray

```

Output:
```
Work functions:  [2.282548 5.012548 4.082548 2.902548]
```

## 2.4 - NumPy and precision (6 points)

Let's use NumPy to explore issues related to floating-point precision some more.

### Part A (2 points)

In the cell below, __use NumPy to calculate__ $\sin(0)$ and $\sin(\pi)$, and print them both out.


new HW idea: numpy + precision, explain why `np.sin(pi)` isn't zero.  Can I use `np.sum` on a long array to show off accumulation of roundoff error...?

pull from hw7 for numpy problems.

<!-- grade_id: cell-b3c5efc0587e8624, points: 0, grade: false, solution: true -->
```python
## YOUR CODE HERE
import numpy as np
from math import sin
import math
pi = math.pi
print (np.sin(pi), np.sin(0))
```

Output:
```
1.2246467991473532e-16 0.0
```

Now based on what you know about floating-point precision, explain __why these two values are different__.  Does the size of the difference match what you learned about the precision of 64-bit floating point numbers?

<!-- grade_id: cell-004c52cde60ab9a6, points: 2, grade: true, solution: true -->
The number π cannot be represented exactly as a floating-point number. So, np.sin(pi) doesn't give me π, it gives me 3.1415926535897931.

And sin(3.1415926535897931) is in fact something like 1.22e-16

### Part B (2 points)

Although the truncation error we saw above was very small, it can still be significant if we're not careful.  In particular, if a truncation error occurs many times, we can end up with __accumulated errors__ that can be significant.

In the cell below, __create an array containing the value $\pi$ 1 million times__.  Then take the sine of the array, and sum over it with `np.sum()`.  _Now_ how large is the difference from zero?

<!-- grade_id: cell-6765b76ea3ef7964, points: 2, grade: true, solution: true -->
```python
import numpy as np
from math import sin
import math
print ('the difference from zero is:', np.sum(np.sin(np.full((1, 1000000), math.pi))))
```

Output:
```
the difference from zero is: 1.2246467991473518e-10
```

### Part C (2 points)

We're insulated from these issues with precision by using 64-bit floating point, but NumPy supports other representations that use small amounts of memory.  Use the `np.array()` function to create another array with 1 million copies of $\pi$, but this time __in 32-bit precision (using data type `np.float32`.)__  You should get a much larger accumulated error!

<!-- grade_id: cell-5b23f603432911f3, points: 2, grade: true, solution: true -->
```python
## YOUR CODE HERE
import numpy as np
piarray = np.ones(1000000,dtype=np.float32)*np.pi
piarray = np.sin(piarray)
print('the error is:', np.sum(piarray))
```

Output:
```
the error is: -0.08742271
```

```python

```