# Homework \#1

## PHYS 2600, Spring 2019

__Important note on entering text:__  Some of the boxes below ask you to enter some text, with the prompt "_Type your answer here using Markdown._"  We haven't learned any Markdown yet - but that's fine!  Regular text is a subset of Markdown, so you should just type plain text into those cells for your answers.

## 1.1 - Introduction (6 points)

At this level of study in undergraduate physics, students tend to have a pretty wide range of experience in computing, particularly using it for scientific work.  Tell us a little about your background below.  First, do you have any experience with Python or other programming languages?  If so, how much?

<!-- grade_id: intro_background, points: 2, grade: true, solution: true -->
__This was an interesting survey!  A little under half the class has a good amount of programming experience; the rest of you have no formal programming training, or some experience with things like Mathematica or Matlab.  Most of you don't have extensive experience with Python specifically, which means you're in the right place!__

Have you taken any computer programming courses before, either here at CU or elsewhere?  If so, what courses did you take?

<!-- grade_id: intro_courses, points: 2, grade: true, solution: true -->
_Type your answer here using Markdown._

After we get through the basics of Python and some standard numerical techniques, we'll have some room to cover more specific examples.  Is there any one subject, physics application or numerical technique, that you'd be especially interested in seeing?  (I may choose from your answers to this question to plan lectures or assignments late in the semester, so this is your chance to guide the class!)

<!-- grade_id: intro_interests, points: 2, grade: true, solution: true -->
__Thanks for your suggestions on this!  Many of the requested topics are already on my list, including orbital dynamics/solar system modeling, the Schrodinger equation, projectile motion, Monte Carlo methods, big data, and (a little) machine learning.  For other topics like controlling hardware or modeling complex systems like the weather, we won't get there, but what you learn in this class will be a good start!__

## 1.2 - Syntax and Semantics (4 points)

Describe in your own words the difference between _syntax_ and _semantics_ in computer programming.

<!-- grade_id: syntax, points: 2, grade: true, solution: true -->
__Syntax is the basic rules for how the pieces of a language can be put together; in Python, syntax tells us how functions and objects can be assembled into valid statements (pieces of code that will do something.)__

__Semantics is the relationship between language and _meaning_.  Human languages are allowed to have ambiguous semantics, which can vary depending on context.  In programming languages, the semantics correspond to data and operations carried out by the computer, so semantics in Python are exact: every program does one thing!__

Computers aren't able to flexibly interpret what we tell them to do; they are always perfectly literal.  This is a challenge if you're trying to build a computer to handle a complicated and dynamical process, for example building a self-driving car.

Look up the turn-by-turn driving directions from Boulder to Denver using [Google Maps](http://maps.google.com/).  Now imagine that a self-driving car has been given these directions and nothing more; what are some ways in which the trip could go horribly wrong, compared to what a human driver would do?  Give an example.

<!-- grade_id: self_driving, points: 2, grade: true, solution: true -->
__My favorite potential mistakes I saw related specifically to the map directions:__

* "Take the exit" might cause the car to try to bring the exit sign with it;
* "Use the right 2 lanes" might make the car use both lanes at once, maybe swerving back and forth;
* "Get on I-25 N from W Colfax Ave" doesn't specify _how_ to get on, which could get interesting!

__Aside from misinterpretations like this, most of you answered regarding the bigger problem, which is that a car has to fill in lots of info not on the directions about traffic laws, dealing with pedestrians, other cars, etc.  Your list of problems is a small subset of what engineers trying to really build these things worry about every day!__

## 1.3 - Planetary Weight Conversion (4 points)

Suppose we are standing on an arbitrary planet of mass $M$ and radius $R$.  The force of gravity on us (mass $m$) from the planet is equal to:

\\[
|\vec{F_g}| = \frac{GmM}{R^2}.
\\]

In the code cell below, __implement a formula to calculate the weight (in Newtons) on another planet__ of a person who weighs 800 N on the Earth.  Use it for two planets: Mars (save your answer as a variable named `weight_Mars`) and Mercury (saved as `weight_Mercury`.)  Their mass and radius are:

* Mars: $R = 0.532 R_E$, $M = 0.107 M_E$
* Mercury: $R = 0.383 R_E$, $M = 0.0553 M_E$

The correct answer for the person's weight on Mars is given in the test cell below, to help you check your work.  (Tests are always a good idea in computer programming, to make sure the program does what you expect!)

<!-- grade_id: planets, points: 0, grade: false, solution: true -->
```python
weight_Earth = 800  # N

### BEGIN SOLUTION
M_Mars = 0.107  # units of M_E, Earth's mass
M_Mercury = 0.0553
R_Mars = 0.532  # units of R_E, Earth's radius
R_Mercury = 0.383

# Force of gravity is proportional to M / R**2, so weight on Earth is rescaled to
# (M/M_E) / (R/R_E)**2.  Working in our units, we multiply by M/R**2.
weight_Mars = weight_Earth * M_Mars / R_Mars**2
weight_Mercury = weight_Earth * M_Mercury / R_Mercury**2

### END SOLUTION

```

<!-- grade_id: planets_tests, points: 4, grade: true, solution: false -->
```python
# Testing cell; don't worry about the details yet, 
# but it should run without error if your answer for Mars in weight_Mars is correct.

import numpy.testing as npt

print("Weight on Earth: %g N" % weight_Earth)
print("Weight on Mars: %d N" % weight_Mars)
print("Weight on Mercury: %d N" % weight_Mercury)

npt.assert_allclose(weight_Mars, 302, rtol=0.01, atol=0)

```

Output:
```
Weight on Earth: 800 N
Weight on Mars: 302 N
Weight on Mercury: 301 N
```

## 1.4 - Clock Math (6 points)

As we discussed in lecture 3, we can use Python's `%` (modulo) operator to do "clock math": adding together times on a 12-hour clock face.  There are more interesting applications of `%` that we'll see later on, so to prepare for those, let's start with some clock-math exercises to get familiar with `%`.

### Part A (2 points)

Let's start with a couple of simple _test cases_ where we know the answer before we calculate anything.

* __Test one__: if we add 12 hours to 3:00, it will still be 3:00.
* __Test two__: if we add 9 hours to 3:00, it will be 12:00.

Note that in Python, if we compute `x % 12`, the result will be between 0 and 11.  So we'll adopt the convention that "12 o'clock" is 0 in the code below.

In the cell below, __use the `%` operator__ and the variable "`now_time`" (set to 3:00) to confirm the two test cases above.  Save the results of adding each amount of hours to `now_time` to the variables `test_time_one` and `test_time_two`.

<!-- grade_id: clock_simple, points: 0, grade: false, solution: true -->
```python
now_time = 3

### BEGIN SOLUTION
test_time_one = (now_time + 12) % 12
test_time_two = (now_time + 9) % 12
### END SOLUTION

```

Now the cell below will run a test to make sure your results for `test_time_one` and `test_time_two` are correct.

(Don't worry about the details of the `assert` command right now; all you need to know is that the cell will run with no complaints if your cell above is correct, and will raise an error otherwise.)

<!-- grade_id: clock_simple_tests, points: 2, grade: true, solution: false -->
```python
# Testing cell; should run with no error message if your answer above is right!
print("test_time_one = %d (should be 3)" % test_time_one)
print("test_time_two = %d (should be 0)" % test_time_two)

assert (test_time_one == 3)
assert (test_time_two == 0)
```

### Part B (2 points)

Now let's answer a question where we don't know the answer.  If the current time is 7:00, what time of day will it be 1 million hours from now?  (Give your answer as an integer from 0 to 11, with 0 corresponding to 12:00; save it as a variable called `million_time`.)

<!-- grade_id: clock, points: 0, grade: false, solution: true -->
```python
now_time = 7

### BEGIN SOLUTION
million_time = (now_time + 1000000) % 12
### END SOLUTION

```

<!-- grade_id: clock_tests, points: 2, grade: true, solution: false -->
```python
# Hidden testing cell: do not delete or you will lose points!

print("A million hours from 7:00, it will be %d:00." % million_time)

```

Output:
```
A million hours from 7:00, it will be 11:00.
```

### Part C (2 points)

Suppose we start at 7:00 AM, to be more specific.  __Is the time you found above (1 million hours from 7:00 AM) AM or PM?__  Use the code cell below to write your solution, and then __give your answer (AM or PM) in a comment.__

<!-- grade_id: clock_am_pm, points: 2, grade: true, solution: true -->
```python
# Distinguishing AM and PM means we have a 24-hour clock, instead of a 12-hour one.
# Let's use 0 to denote 12 midnight, so 7 is 7:00 AM.
# 
# Now we use modulo again, but this time with a base of 24:
million_time_AM_PM = (7 + 1000000) % 24
print(million_time_AM_PM)

# The code prints '23' for million_time_AM_PM, which is __11:00 PM.__
```

Output:
```
23
```

## 1.5 - Mistakes Were Made (9 points)

All three of the following code examples are implemented incorrectly, containing one or more errors that prevent them from running.  __Fix the code in each example__ to run the correct formula, and give the right answer.

(_Note:_ I emphasize that you have to run the correct formula in full - no credit for just entering `print(right answer)`!)

### Part A (3 points)

__Q:__ What is the kinetic energy of a 1500 kg car moving at 12 m/s?

__Formula:__ $KE = \frac{1}{2} m v^2$

__A:__ We should get 108,000 J as the answer.

Copy and paste into the code cell below, then fix the mistakes.

```python
kinetic-energy = (1/2) * 1500 * v^2

print(kinetic-energy)

```

<!-- grade_id: cell-057f944b5694a43a, points: 3, grade: true, solution: true -->
```python
v = 12  # Mistake 1: v wasn't defined.

# Mistake 2: "kinetic-energy" is an invalid variable name
# Mistake 3: ^ isn't the power operator, ** is.

kinetic_energy = (1/2) * 1500 * v**2

print(kinetic_energy)
```

Output:
```
108000.0
```

### Part B (3 points)

__Q:__ An air hockey puck is measured at coordinates x=30 cm, y=30 cm.  How far is it from the origin?

__Formula:__ $d = \sqrt{x^2 + y^2}$

__A:__ We should get 42.4 cm.

Copy and paste into the code cell below, then fix the mistakes.

```python
x = 30
    y = 30

    d = x**2 + y**2**(1//2)

print(d)
```

<!-- grade_id: mistakes_2, points: 3, grade: true, solution: true -->
```python
x = 30

# Mistake 1: can't have extra whitespace at start of lines!
y = 30

# Mistake 2: 1//2 is integer division, we want 1/2 or 0.5 for square root.
# Mistake 3: Need parentheses around (x**2 + y**2) to take the square root of the sum.
d = (x**2 + y**2)**(1/2)

print(d)

```

Output:
```
42.42640687119285
```

### Part C (3 points)

__Q:__ A block weighing 800 grams sits on a surface with coefficient of static friction mu = 0.22.  How much force is required to overcome static friction?  (Take g = 9.8 m/s^2.)

__Formula:__ $F_f = \mu m g$

__A:__ The force required is about 1.72 N.

Copy and paste into the code cell below, then fix the mistakes.

```python
mu = '0.22'
F_N = 800 * 9.8
F_f = mu * F_N

print F_f
```

<!-- grade_id: mistakes_3, points: 3, grade: true, solution: true -->
```python
# Mistake 1: mu should be a number, not a string.
mu = 0.22

# Mistake 2: working in MKS, a mass of 800 grams = 0.8 kg should be "0.8"
F_N = 0.8 * 9.8
F_f = mu * F_N

# Mistake 3: print was missing the parentheses () required in Python 3.
print(F_f)
```

Output:
```
1.7248
```