# Homework #7

## PHYS 2600, Spring 2019

__Important notice:__ All cells in your notebook will be run, start to finish, using a fresh kernel when it is graded!  To make sure the graded notebook looks like what you expect, we recommend restarting the kernel and selecting "Cells > Run All" in the menu above before you submit.

If you worked collaboratively on this assignment, __include the names of your collaborators in the cell below:__

## 7.0 - Homework correction (3 points)

In the cell below, explain and correct __one mistake__ you made on your _last graded_ homework assignment.  If you got a perfect score, great, tell us what your favorite problem was, or use the space to give any other feedback you might have on the class/tutorials/homework.

<!-- grade_id: cell-1689332322e959be, points: 3, grade: true, solution: true -->
_Type your answer here using Markdown._

## 7.1 - More fun with lists (10 points)

__Important instruction:__ Your solution to all of the problems below _must_ use list operations to manipulate the starting lists or tuples you are given, in order to produce the answer.  Just entering the answer by hand as a new list is worth no credit!

Run the cell below to initialize some lists and tuples that you'll have to manipulate in your answers.

```python
# Initialization cell - run me first!

SI_smaller_prefixes = ('milli','micro','nano','pico')
SI_bigger_prefixes = ('kilo', 'mega', 'giga', 'tera')

material_names = ['water', 'selenium', 'quartz', 'potassium']
melting_points = [0, 220, 1470, 63]
boiling_points = [100, 688, 2230, 762]

some_points_xyz = [
    (3, 1, 1),
    (2, 1, -2),
    (-1, 0, -1),
    (-7, -3, -4),
    (5, -5, 2),
    (-3, 2, 0),
]
```

### Part A (3 points)

__Combine the tuples__ `SI_smaller_prefixes` and `SI_bigger_prefixes` into a single tuple `SI_prefixes`.  Make sure your combined list is __in order, from smallest to largest prefix (i.e. pico to tera).__

_(Hint: you could use a `for` loop here to build the tuple you want, more likely two `for` loops to traverse the two different lists.  Or you can skip the loop and just add the two prefix tuples up, but then they'll be in the wrong order...slicing with a step size of $-1$ can let you reverse a list or tuple.)_

<!-- grade_id: cell-1ac29bbb3bcad33f, points: 0, grade: false, solution: true -->
```python

### BEGIN SOLUTION
# Reverse and add together
SI_prefixes = SI_smaller_prefixes[::-1] + SI_bigger_prefixes

# Or, use a loop:
SI_prefixes = ()
for i in range(len(SI_smaller_prefixes)):
    SI_prefixes += (SI_smaller_prefixes[-(i+1)],)
for i in range(len(SI_bigger_prefixes)):
    SI_prefixes += (SI_bigger_prefixes[i],)

SI_prefixes
### END SOLUTION


```

Output:
```
('pico', 'nano', 'micro', 'milli', 'kilo', 'mega', 'giga', 'tera')
```

<!-- grade_id: cell-e74f47c5af91ffac, points: 3, grade: true, solution: false -->
```python
assert SI_prefixes == ('pico', 'nano', 'micro', 'milli', 'kilo', 'mega', 'giga', 'tera')
```

### Part B (3 points)

Three related lists above contain the names of some materials, their melting points (in degrees C), and their boiling points (also degrees C).  Combine them to __create a single list__ `material_melt_boil_points`, which consists of  __tuples of the form:__ `(material name, melting point, boiling point)`.

<!-- grade_id: cell-f4e474f8024f342b, points: 0, grade: false, solution: true -->
```python

### BEGIN SOLUTION
material_melt_boil_points = []
for i in range(4):
    material_melt_boil_points.append( (material_names[i], melting_points[i], boiling_points[i]) )

material_melt_boil_points
### END SOLUTION

```

Output:
```
[('water', 0, 100),
 ('selenium', 220, 688),
 ('quartz', 1470, 2230),
 ('potassium', 63, 762)]
```

<!-- grade_id: cell-3c6ccf0838796781, points: 3, grade: true, solution: false -->
```python
solution_materials_list = [
    ('water', 0, 100),
    ('selenium', 220, 688),
    ('quartz', 1470, 2230),
    ('potassium', 63, 762),
]
assert material_melt_boil_points == solution_materials_list
```

### Part C (4 points)

Starting with the list `some_points_xyz`, which is a list of Cartesian points $(x,y,z)$, produce a list `some_points_cyl` which consists of __three-element tuples__ $(r, \phi, z)$, corresponding to the same points in cylindrical coordinates - defined by:

\\[
r \equiv \sqrt{x^2 + y^2}, \\
\phi = \tan^{-1} (y/x),
\\]
and $z$ is unchanged.

_(Note: be careful with how you compute the angle `phi` - check back on tutorial 15 for the difference between the `arctan` and `arctan2` functions.)_

<!-- grade_id: cell-2f71c1b87bbb2aa8, points: 0, grade: false, solution: true -->
```python
import numpy as np


### BEGIN SOLUTION
some_points_cyl = []
for x, y, z in some_points_xyz:    
    r = np.sqrt(x**2 + y**2)
    phi = np.arctan2(y, x)
    
    some_points_cyl.append( (r, phi, z) )
    
some_points_cyl
### END SOLUTION

```

Output:
```
[(3.1622776601683795, 0.3217505543966422, 1),
 (2.23606797749979, 0.4636476090008061, -2),
 (1.0, 3.141592653589793, -1),
 (7.615773105863909, -2.7367008673047097, -4),
 (7.0710678118654755, -0.7853981633974483, 2),
 (3.605551275463989, 2.5535900500422257, 0)]
```

<!-- grade_id: cell-bd3677ed30d02a01, points: 4, grade: true, solution: false -->
```python
import numpy.testing as npt

test_points = [
    (3.16, 0.32, 1),
    (2.24, 0.46, -2),
    (1.00, 3.14, -1),
    (7.62, -2.74, -4),
    (7.07, -0.79, 2),
    (3.61, 2.55, 0),
]

npt.assert_allclose(test_points, some_points_cyl, atol=1e-2)
```

## 7.2 - Friction and lists (12 points)

### Part A (4 points)

Suppose we're performing an experiment in 1-d kinematics with friction.  We have a collection of blocks, initially all  with the same coefficients of friction $\mu_S = \mu_K = 0.2$.  (This is a bit unrealistic since $\mu_S > \mu_K$ generally, although for some materials $\mu_S = \mu_K$ isn't a terrible approximation.)

The blocks have various masses, and we apply different forces to them.  We'd like a function that will check whether the blocks will move or not, and compute the acceleration only if they do move.

I have started filling in the function `block_accel_friction` below, including a docstring.  __Implement the missing logic__ before the `return` statement, so that the function will behave as described in the docstring.



<!-- grade_id: cell-86e4e4d36c4d094c, points: 0, grade: false, solution: true -->
```python

def block_accel_friction(m, F, mu=0.2):
    """
    Checks whether a block will move, overcoming
    static friction, and returns the acceleration if so.

    Arguments:
    ==========
    m: mass of block  [kg]
    F: applied linear force  [N]
    mu: coefficient of friction (default: 0.2.)
    
    Returns:
    ========
    A tuple (block_will_move, accel) = (<bool>, <object>), 
    depending on whether the block moves:
    
    * Returns (False, None) if the block does not move;
    * Returns (True, accel) if the block does move.

    Zero acceleration (F = F_friction) is considered *not* moving.
    """
    
    g = 9.8
    F_friction = mu * m * g

    ### BEGIN SOLUTION
    
    if F <= F_friction:
        block_will_move = False
        accel = None
    else:
        block_will_move = True
        accel = (F - F_friction) / m

    ### END SOLUTION
    

    return (block_will_move, accel)
    
```

<!-- grade_id: cell-3cc96e53e7e2a822, points: 4, grade: true, solution: false -->
```python
import numpy.testing as npt

npt.assert_allclose(block_accel_friction(1.0, 4.0), (True, 2.04))
assert block_accel_friction(4.0, 4.0) == (False, None)
assert block_accel_friction(1.0, 4.0, mu=0.5) == (False, None)
npt.assert_allclose(block_accel_friction(4.0, 20.0), (True, 3.04))
```

### Part B (4 points)

Below we have two tuples about our experiment, now with some blocks of various $\mu$.  `block_list` consists of 3-element tuples `(mass, force, mu)` for a single block.  `block_color_list` is a list of the same length containing the color each block is painted (which is how we keep track of them in the experiment.)

Use a `for` loop to __produce a single list called `blocks_that_move`__ made up of tuples `(color, accel)`, __only for those blocks which overcome static friction and move.__

_(Hint: if you're not sure how to do this at first, can you start by just creating the list `(color, accel)` for every single block?  If you called_ `block_accel_friction` _to get `accel`, what did you do with the boolean_ `block_will_move` _that also gets returned?  Can you use that to reduce to the list you want?)_

<!-- grade_id: cell-a92ed3a8297329db, points: 0, grade: false, solution: true -->
```python
# block_list contains data as (mass, force, mu)
block_list = (  
    (0.5, 0.8, 0.2),
    (1.0, 3.0, 0.2),
    (2.0, 4.4, 0.2),
    (3.0, 5.6, 0.1),
    (2.0, 2.2, 0.3),
    (2.0, 1.6, 0.1),
    (1.0, 2.6, 0.1),
)
block_color_list = ('red', 'orange', 'yellow', 'green', 'blue', 'white', 'black')

blocks_that_move = []

### BEGIN SOLUTION
for i, (m, F, mu) in enumerate(block_list):
    block_will_move, accel = block_accel_friction(m, F, mu=mu)
    if block_will_move:
        blocks_that_move.append((accel, block_color_list[i]))

### END SOLUTION

blocks_that_move  # Correct solution should contain four blocks out of seven.
```

Output:
```
[(1.0399999999999998, 'orange'),
 (0.24, 'yellow'),
 (0.8866666666666663, 'green'),
 (1.62, 'black')]
```

<!-- grade_id: cell-d6989a2db387c49a, points: 4, grade: true, solution: false -->
```python
import numpy.testing as npt

assert len(blocks_that_move) == 4

```

### Part C (4 points)

Finally, suppose we're thinking about designing the next iteration of the experiment, and we have a variety of block masses and applied forces to combine, given in the lists `mass_list` and `force_list`.

__Produce a list `mass_force_combos` containing all combinations of any mass and any force such that the block will move.__  (Assume the default value of friction $\mu = 0.2$.)

_(Hint: you'll need to start with a list of __all__ possible combinations of mass and force from the two lists, and then filter it down with the function from part A.  The list of all pairwise combinations from two lists is known as the __direct product__ or __outer product__.  Nesting two `for` loops inside each other is a good way to start.)_

<!-- grade_id: cell-df73c09a42f8445b, points: 0, grade: false, solution: true -->
```python
mass_list = (0.5, 1.0, 2.0, 3.0,)
force_list = (0.8, 2.0, 3.4, 5.6)

mass_force_combos = []
### BEGIN SOLUTION
for m in mass_list:
    for F in force_list:
        # `_` is a common convention in Python for an ignored placeholder variable.
        block_will_move, _ = block_accel_friction(m, F)
        if block_will_move:
            mass_force_combos.append((m,F))
### END SOLUTION


mass_force_combos  # Correct solution has 7 different (m, F) combinations.
```

Output:
```
[(0.5, 2.0),
 (0.5, 3.4),
 (0.5, 5.6),
 (1.0, 2.0),
 (1.0, 3.4),
 (1.0, 5.6),
 (2.0, 5.6)]
```

<!-- grade_id: cell-42208ee68300e407, points: 4, grade: true, solution: false -->
```python
assert len(mass_force_combos) == 7

```

## 7.3 - A taste of cryptography (16 points)

Now that we've met lists and tuples, we have a new way to look at our old friend the _string_.  A string is a __sequence__, just like a list or tuple: it is an ordered set of objects.  In the case of a string, the objects it's made of must be __characters__ - single letters, numbers, punctuation marks, etc.

Like a tuple, a string is an _immutable_ sequence, which means we can't use things like `append` or use `=` to overwrite individual characters.  However, basically all of the methods we've seen for tuples also work as expected on strings, including indexing, slicing, and `len`.  We can even use a `for` loop to iterate over a string, which gives us one character at a time:

```python
s = 'abcxyz'
print(s[1])
print(len(s))
print(s[3:6])

print("For loop:")
for char in s:
    print(char)
```

Output:
```
b
6
xyz
For loop:
a
b
c
x
y
z
```

The most immediate and interesting application of this sort of character-level string manipulation is __cryptography__, or secure communications.  (The textbook, Langtangen, has another example - working with DNA sequences - at the end of chapter 3.)

The idea is simple: if we can devise a secret algorithm to __encode__ a message by converting it from plain text into something unreadable, and if the algorithm is reversible so we can __decode__ the message back to plain text, then we have a way to send messages to another person so that any third party who doesn't know the algorithm can't intercept and read them.

For simplicity, let's focus only on messages in English consisting of lower-case letters.  (It's not any harder to do mixed upper-case and lower-case letters, just more complicated.)  One of the simplest classes of cryptographic algorithms are the __ROT__, or "rotation", algorithms.  The __ROT(n)__ algorithm takes each letter in a message, and "rotates" it around the alphabet by n places.  

So, for example, ROT(1) would map a --> b, b --> c, and so on (with the end wrapping around, so z --> a.)  The message `"hello"` becomes `"jfkkp"`; to 'decode' back to `"hello"` we just reverse the map.  The wrapping-around at the end is like counting on a clock (an hour after 12 is 1), so this is another example of _modular arithmetic_, as we encountered back in lecture 2.




### Part A (6 points)

To implement a ROT algorithm, the first thing we need is a mapping between letters and their corresponding numbers.  Since letters are really stored as binary numbers underneath, we can just use the numeric representation that Python already has!

Given a single character `c`, the built-in function `ord(c)` converts it into the underlying number in the encoding scheme we are using (in this case, ASCII, which we discussed briefly back in lecture 3.)  The `chr()` function goes back the other way: for example,

```python
print(chr(97))
print(ord('a'))
print(ord('z'))
```

Output:
```
a
97
122
```

As these results imply, the codes given by `ord()` for lower-case letters run continuously from 97 to 122.  However, this is a little annoying, since the modular arithmetic we'll need is much more natural to handle if our codes range from 0 to 25 instead.

We also need to think about how to handle things that _aren't_ lower-case letters.  Most messages will be harder to read without punctuation, and _much_ harder to read without spaces!  Keep this in mind as we're working towards implementing the full algorithm.  (Of course, space and punctuation also make it easier to guess the true message - so the professionals tend to just not use them.)

To make implementing the ROT(n) algorithm easier, __define three functions below__:

* `is_letter(c)` should accept a single character (length-one string), and return `True` if it is a lower-case letter and `False` otherwise.  _(Remember, `ord(c)` for any character will give you a code, and we just worked out which codes are letters...)_
* `char_to_num(c)` should also accept a single character.  If `c` is _not_ a lower-case letter, it should just return `c` back; otherwise, it should return 0 for `a`, 1 for `b`, and so on.  _(You should use_ `is_letter(c)` _here!)_
* `num_to_char(n)` is the inverse of `char_to_num`: it should accept a number from 0-25 and return the corresponding lower-case letter.

<!-- grade_id: cell-b2bb6a0aa2db1960, points: 0, grade: false, solution: true -->
```python
def is_letter(c):    
    ### BEGIN SOLUTION
    return 97 <= ord(c) <= 122
    ### END SOLUTION

def char_to_num(c):
    ### BEGIN SOLUTION
    if not is_letter(c):
        return c
    
    return ord(c) - 97
    ### END SOLUTION
    
def num_to_char(n):
    ### BEGIN SOLUTION
    if not type(n) == int:
        return n
    
    return chr(n+97)
    ### END SOLUTION
```

<!-- grade_id: cell-18f8dc8bff4a199b, points: 6, grade: true, solution: false -->
```python
assert is_letter('q')
assert not is_letter('3')
assert not is_letter('X')

assert char_to_num('z') == 25
assert char_to_num('n') == 13
assert char_to_num(' ') == ' '
assert char_to_num('A') == 'A'

assert num_to_char(4) == 'e'
assert num_to_char(23) == 'x'

assert num_to_char(char_to_num('k')) == 'k'
assert char_to_num(num_to_char(13)) == 13
```

### Part B (6 points)

Now __implement the ROT(n) algorithm__ in the function `encode_rot_n(message, n)` below.  Your function should return a string in which every lower-case letter is replaced by its 'rotated' equivalent, according to the value of `n` given as an argument.

A good test case to consider is ROT(13), which is particularly simple: the first half of the alphabet gets mapped to the second half, and vice versa.  So `'a'` --> `'m'` and `'m'` --> `'a`', and so on.

_(Hint: remember,_ `char_to_num` _returns either an `int` type or a `str` type, and you only want to apply the "rotation" if it gives you an `int` back!  A quick type check of the form `if type(...) == int:` should let you discriminate between lower-case letters and everything else in your code.)_

<!-- grade_id: cell-7acbba8650cb90fd, points: 0, grade: false, solution: true -->
```python
def encode_rot_n(message, n):
    ### BEGIN SOLUTION
    new_message = ''
    
    for char in message:
        cnum = char_to_num(char)
        if type(cnum) == int:        
            rot_cnum = (cnum + n) % 26
            rot_char = num_to_char(rot_cnum)
        else:
            rot_char = char

        new_message += rot_char
    
    return new_message
    ### END SOLUTION
    
```

<!-- grade_id: cell-9f1970ba14f6b42d, points: 6, grade: true, solution: false -->
```python
assert encode_rot_n('colorado', 19) == 'vhehktwh'
assert encode_rot_n('hello world', 13) == 'uryyb jbeyq'
assert encode_rot_n('Testing upper case', 4) == 'Tiwxmrk yttiv gewi'
assert encode_rot_n(encode_rot_n('back and forth!', 13), 13) == 'back and forth!'

```

We only wrote a ROT(n) encoder, but you're now fully equipped to _decode_ messages as well!  The reason is simple: if we combine ROT(n) with ROT(26-n), then every letter in the resulting twice-encoded message has been rotated by 26 places - which means back to where it started!

```python
encode_rot_n('yvccf nficu', 26-17)  # 'yvccf nficu' is 'hello world' in ROT(17)
```

Output:
```
'hello world'
```

### Part C (4 points)

Let's test out your codebreaking skills.  __The instructions for part C of this problem are encoded in ROT(n) below.__ (Which $n$?  If I told you, it would be too easy...)

_(Hint: if your decoded message looks messy and doesn't have any line breaks, use `print()` on it to get it to print out nicely.)_

<!-- grade_id: cell-fd7fd94c730602e8, points: 0, grade: true, solution: true -->
```python
instr_encoded = """
fchohwcb ozucfwhvag ofs bch jsfm gsqifs.  wt ks gigdsqh o fchohwcb
ozucfwhva, ks cbzm bssr hc hsgh 26 jsfgwcbg ct wh hc rsqcrs hvs asggous.
sjsb kcfgs, wh'g sogm hc uisgg kvoh b wg tfca qcaacb gwbuzs-zshhsf kcfrg
zwys "w" obr "o" - mci aom vojs rcbs hvoh hc rsqcrs hvwg, wb toqh.
qfmdhcufodvm wg cbzm og uccr og hvs ozucfwhva ks igs!

ks qob joghzm wadfcjs hvs gsqifwhm ct fch b pm igwbu o cbs-hwas dor.
cbs-hwas dorg ofs o gseisbqs ct fobrca wbtcfaohwcb kvwqv wg gvofsr pshkssb
hkc dscdzs kvc kobh hc gsbr obr fsqswjs o asggous.  hvwg hsqvbweis wg
ibpfsoyopzs wt hvs cbs-hwas dor whgszt wg hfizm fobrca, pih wh rcsg fseiwfs
hvs gsbrsf obr fsqswjsf hc assh obr slqvobus hvs dor pstcfsvobr.

wadzsasbh hvs tibqhwcb sbqcrs_fch_b_chd pszck.  hvs jofwopzs cbs_hwas_dor gvcizr
ps o zcbu ghfwbu ct zcksf-qogs zshhsfg, kvwqv ks vojs ozfsorm aoddsr hc biapsfg
tfca 0 hc 25.  gc hvs twfgh qvofoqhsf wb asggous gvcizr ps sbqcrsr kwhv sbqcrs_fch_b,
obr b seioz hc qvof_hc_bia() ct hvs twfgh qvofoqhsf wb cbs_hwas_dor, obr gc cb.

sbqcrs_fch_b_chd ozgc wbqzirsg o pcczsob ofuiasbh, rsqcrs.  wt rsqcrs wg hfis,
hvsb mci gvcizr ghwzz fsor b tfca hvs cbs_hwas_dor, pih hvsb fchohs pm 26 - b.
wt mci'js rcbs sjsfmhvwbu qcffsqhzm, mci'zz ps opzs hc sbqcrs obr rsqcrs, obr
fsqcjsf hvs asggous mci ghofhsr kwhv.

ks vojs hc rsqwrs vck hc rsoz kwhv bcb-zshhsfg, cf hvs rsqcrwbu aom towz othsf
hvs twfgh gdoqs oddsofg.  am wadzsasbhohwcb "igsg id" o zshhsf wb hvs cbs-hwas dor
kvsb o bcb-zshhsf wg tcibr.  wb chvsf kcfrg: wt hvs hvwfr qvofoqhsf wb hvs asggous wg
o gdoqs, hvsb hvs hvwfr qvofoqhsf wb hvs cbs-hwas dor wg igsr id, obr hvs bslh qvofoqhsf
wg sbqcrsr kwhv hvs tcifhv tfca hvs cbs-hwas dor.  hvwg wg o pwh koghstiz, pih wh'g
ozgc o zwhhzs sogwsf hc wadzsasbh.  ouowb, wt mci rc hvwg kfcbu, mci'zz ybck psqoigs
mcif rsqcrwbu kwzz dfcriqs cbs kcfr obr hvsb uofpous othsf hvs gdoqs!

hvsfs ofs bc hsghg tcf hvwg dofh!  wbghsor, mci aigh obgksf wb hvs twboz
aofyrckb qszz pszck: kvoh wg hvs gsqfsh qcrs kcfr sbqcrsr wb am zogh asggous?

vsfs wg mcif cbs hwas dor hc rsqcrs hvs zogh asggous pszck:

lmavmpxxmkknbqhkqyzjyamdzlofucyqeqfasasbheehemeqwyfkispvyxunphjydzbneemyfrrwogxqxuakmkkpulvesakzczxwqbiz
"""

### BEGIN SOLUTION
# The message below will be revealed if we decode the above using
# rot-14 (equivalently, encode with rot-12.)

instr_decoded = """
rotation algorithms are not very secure.  if we suspect a rotation
algorithm, we only need to test 26 versions of it to decode the message.
even worse, it's easy to guess what n is from common single-letter words
like "i" and "a" - you may have done that to decode this, in fact.
cryptography is only as good as the algorithm we use!

we can vastly improve the security of rot n by using a one-time pad.
one-time pads are a sequence of random information which is shared between
two people who want to send and receive a message.  this technique is
unbreakable if the one-time pad itself is truly random, but it does require
the sender and receiver to meet and exchange the pad beforehand.

implement the function encode_rot_n_otp below.  the variable one_time_pad should
be a long string of lower-case letters, which we have already mapped to numbers
from 0 to 25.  so the first character in message should be encoded with encode_rot_n,
and n equal to char_to_num() of the first character in one_time_pad, and so on.

encode_rot_n_otp also includes a boolean argument, decode.  if decode is true,
then you should still read n from the one_time_pad, but then rotate by 26 - n.
if you've done everything correctly, you'll be able to encode and decode, and
recover the message you started with.

we have to decide how to deal with non-letters, or the decoding may fail after
the first space appears.  my implementation "uses up" a letter in the one-time pad
when a non-letter is found.  in other words: if the third character in the message is
a space, then the third character in the one-time pad is used up, and the next character
is encoded with the fourth from the one-time pad.  this is a bit wasteful, but it's
also a little easier to implement.  again, if you do this wrong, you'll know because
your decoding will produce one word and then garbage after the space!

there are no tests for this part!  instead, you must answer in the final
markdown cell below: what is the secret code word encoded in my last message?

here is your one time pad to decode the last message below:

xymhybjjywwznctwcklvkmyplxargokcqcrmementqqtqyqcikrwuebhkjgzbtvkplnzqqykrddiasjcjgmwywwbgxhqemwloljicnul
"""
### END SOLUTION
```

<!-- grade_id: cell-bd2ec88c24027a3c, points: 0, grade: true, solution: true -->
```python
def encode_rot_n_otp(message, one_time_pad, decode=False):
    ### BEGIN SOLUTION
    new_message = ''
    
    for i in range(len(message)):
        otp_n = char_to_num(one_time_pad[i])
        if decode:
            otp_n = 26 - otp_n
        
        new_message += encode_rot_n(message[i], otp_n)
        
        
    return new_message
    ### END SOLUTION
    
```

```python
final_message = 'qfuz nnbqwcd ybhn dzvr-spptiaqd yp fiz fxsegtq...tlon lpnmomh.  isr iuabvw eojm ry "oqwozrt".'
```

<!-- grade_id: cell-0bfe90a090479781, points: 0, grade: true, solution: true -->
```python
# Scratch space for you to figure this out...

### BEGIN SOLUTION
otp_str = 'xymhybjjywwznctwcklvkmyplxargokcqcrmementqqtqyqcikrwuebhkjgzbtvkplnzqqykrddiasjcjgmwywwbgxhqemwloljicnul'
encode_rot_n_otp(final_message, otp_str, decode=True)
### END SOLUTION

```

Output:
```
'this message will self-destruct in ten seconds...just kidding.  the secret word is "quantum".'
```

<!-- grade_id: cell-9fac702ab525df82, points: 4, grade: true, solution: true -->
_Type your answer here using Markdown._