# Tutorial 11: Flow Control

## PHYS 2600, Spring 2019

```python
# Import cell

%matplotlib inline
import numpy as np
import matplotlib.pyplot as plt
```

## T11.X - FizzBuzz!

_(Special note: this first problem is a __worked example__, which we'll go through together in class.  You are encouraged to fill this in as you follow along, but you won't be graded on whether you've completed it or not.)_

Let's start with a little game, "Fizz Buzz":

https://en.wikipedia.org/wiki/Fizz_buzz

The rules to the game are simple: we count up from 1 out loud, but every time we reach a number divisible by 3, we say "Fizz" instead.  Likewise, every number divisible by 5 is replaced with "Buzz".  If a number is divisible by both 3 and 5, we say both words (so "FizzBuzz").

This particular game is infamous in computer programming circles, because writing a program that can mimic FizzBuzz is a common interview question used to screen entry-level programming jobs.  (So you'll all be prepared for that after this tutorial!)

Our version of the problem: __write a Python program that prints the numbers from 1 to 30, replacing them with "Fizz", "Buzz" or "FizzBuzz" according to the rules above.__  Clearly, this calls for a `while` loop!  Let's write a solution.

<!-- grade_id: cell-cf7015a7552343e3, points: 0, grade: true, solution: true -->
```python

### BEGIN SOLUTION
n = 1
while n <= 30:
    # Build a string, s, to print out
    s = ''

    # Add 'Fizz' to the string if n is a multiple of 3
    if (n % 3) == 0:
        s += 'Fizz'
        
    # Add 'Buzz' to the string if n is a multiple of 5
    if (n % 5) == 0:
        s += 'Buzz'
    
    if s == '':
        # If s is still empty, just print the number
        print(n)
    else:
        # Otherwise, print the string we built
        print(s)

    n += 1
    
### END SOLUTION
```

Output:
```
1
2
Fizz
4
Buzz
Fizz
7
8
Fizz
Buzz
11
Fizz
13
14
FizzBuzz
16
17
Fizz
19
Buzz
Fizz
22
23
Fizz
Buzz
26
Fizz
28
29
FizzBuzz
```

## T11.1 - Branching and looping

### Part A

Use conditionals (`if`/`else` and possibly `elif`) to implement the sign function `sgn(x)`, which is defined as follows:
\\[
{\rm sgn}(x) = \begin{cases} +1&,\ x>0; \\ -1&,\ x<0. \end{cases}
\\]
Use an `assert` statement to make your function raise an error message if it's given the value $x=0$.  (We might prefer for it to return `0` instead of giving an error; this depends on what application we're going to use it for, really.)

<!-- grade_id: cell-b846cddf7de2c4d2, points: 2, grade: true, solution: true -->
```python
def sgn(x):
    ### BEGIN SOLUTION
    
    # Version 1: assert first
    assert x != 0
    if x < 0:
        return -1
    else:
        return 1
    
    # Version 2: use elif
    if x < 0:
        return -1
    elif x > 0:
        return 1
    else:
        assert False
    
    ### END SOLUTION
    

```

<!-- grade_id: cell-79b72f283a6f7fc5, points: 0, grade: true, solution: false -->
```python
## Testing cell

import numpy.testing as npt

npt.assert_equal(sgn(5), 1)
npt.assert_equal(sgn(-3.3), -1)
npt.assert_equal(sgn(0.3), 1)
npt.assert_raises(AssertionError, sgn, 0)
```

### Part B

Python has built-in support for __complex numbers__, under the `complex` type.  (You can think of a `complex` as a pair of `float`s, one for the real part and one for the imaginary part.)  Python uses the engineering convention of writing the imaginary unit as `j`.  __Run the cell below__ to see some quick examples:

```python
x = 2+0j
y = 1j    # Must write "1j" to be recognized as complex; j alone doesn't work.
z = 1+1j

print(x+y)
print(y*y)
print(y.real)
print(y.imag)
print(abs(z))  # "abs" acts as the complex modulus |z| = sqrt(Re(z)**2 + Im(z)**2)
```

Output:
```
(2+1j)
(-1+0j)
0.0
1.0
1.4142135623730951
```

Once we're using complex numbers, there should be no numerical problem with taking the square root of a negative number.  However, by default the `math.sqrt()` and `np.sqrt()` functions will give an error message if we give them a negative input (because for many applications we _don't_ want complex answers.)

There are existing modules for dealing with complex math, of course: `numpy` will do it if you explicitly use a complex `dtype` on your array, and Python has the `cmath` ("complex math") module with complex-aware functions like `cmath.sqrt()`.  But why not just do it ourselves?

__Implement the function `real_sqrt` below__, which should take a single real number `x` as input and return its (potentially complex) square root as follows:

- If `x` is positive, return the square root of `x`, plus the constant `0j`.  (This makes the output `complex` type, for consistency.)
- If `x` is negative, return the square root of `-x`, then multiply it by `1j`.  (The syntax of complex numbers means we have to write it as `1j` and not just `j`.)

<!-- grade_id: cell-0506d1ce981a3802, points: 0, grade: false, solution: true -->
```python
def real_sqrt(x):
    ### BEGIN SOLUTION
    if x > 0:
        return np.sqrt(x) + 0j
    else:
        return np.sqrt(-x) * 1j
    ### END SOLUTION
    
```

<!-- grade_id: cell-0ff4bab315a830bc, points: 0, grade: true, solution: false -->
```python
print(real_sqrt(4))  # should print '2+0j'
print(real_sqrt(-4)) # should print '2j'
print(real_sqrt(0))  # should print '0j'

assert (real_sqrt(-11)**2 - (-11)) <= 1e-6
assert (real_sqrt(45)**2 - 45) <= 1e-6
```

Output:
```
(2+0j)
2j
0j
```

### Part C

Use a `while` loop to implement the function `sum_cubes(n)` below, which should return the sum of the cube of every number from 1 to n __including `n`__, i.e.

$$
S_3(n) \equiv \sum_{i=1}^n i^3 = 1^3 + 2^3 + ... + n^3
$$

_(Hint: your `while` loop will give you one value of `i` at a time, but you want to end up with the sum of __all__ of them.  This is a great opportunity to use an __accumulator variable__: in addition to the variable `i`, you should declare another variable __before the loop__ that will hold the sum of all numbers so far.)_

<!-- grade_id: cell-6a77276ba2492ed2, points: 0, grade: false, solution: true -->
```python
def sum_cubes(n):
    ### BEGIN SOLUTION
    total = 0
    i = 1
    
    while i <= n:
        total += i**3
        i += 1
        
    return total
    ### END SOLUTION
    
```

<!-- grade_id: cell-4b9078242ca7b6ba, points: 0, grade: true, solution: false -->
```python
print(sum_cubes(2))

assert sum_cubes(1) == 1
assert sum_cubes(2) == 9
assert sum_cubes(37) == (37*38/2)**2    # Analytic answer: the sum should give (n(n+1)/2)**2.
```

Output:
```
9
```

### Part D

We can also use conditional tests on NumPy arrays to control program flow.  However, we need to do things a bit differently because arrays have multiple values.  For example, suppose we want to take the square root of an array, but _only if_ it has no negative values. 

__Run the next cell__ to see an example that _will not work_:

```python
x = np.array([-2, -1, 4])
if x > 0:
    print(np.sqrt(x))
```

Output:
```
ValueError                                Traceback (most recent call last)
<ipython-input-15-92ea9616f059> in <module>
      1 x = np.array([-2, -1, 4])
----> 2 if x > 0:
      3     print(np.sqrt(x))

ValueError: The truth value of an array with more than one element is ambiguous. Use a.any() or a.all()
```

You should see a `ValueError` that states "the truth value of an array...is ambiguous."  For once, this is an error message that makes a lot of sense!  __We can't branch on the statement `x > 0`__, because the result is neither `True` nor `False`, it's the array `[False, False, True]`.

The second part of the error message also tells us how to fix this: we use the `.any()` and `.all()` methods.  These take a Boolean test, apply it to an entire array, and then return a single summary `True` or `False` outcome:

```python
print(np.any(x > 0))
print(np.all(x > 0))

# x = [3,4,5]  ## Uncomment me to see the other branch of the if statement below run.

if np.all(x>0):
    print(np.sqrt(x))
else:
    print("Some negative numbers in x, skipping sqrt!")
```

Output:
```
True
False
Some negative numbers in x, skipping sqrt!
```

In terms of a mask array like `x > 0`, the way these two functions work is:

* If __any entry__ in mask `m` is `True`, then `np.any(m)` is `True`; otherwise, it's `False`.
* If __every entry__ in mask `m` is `True`, then `np.all(m)` is `True`; otherwise, it's `False`.

You can see that these are sort of complementary to each other: in the code above, we did the test `np.all(x>0)` to see if our array `x` was entirely positive.  But we could also have done the test `np.any(x<0)`, and the code would work exactly the same.

Now, the exercise.  In the cell below, __implement the function `rescale_array(x)`__, which should take an array `x` and return the same array _repeatedly divided by two_ until all of the entries are between -1 and 1.  (Since arrays are mutable, in this case we begin by making a _copy_ of the array `y`, instead of mutating the original `x`.)

For example, the array `[-3,1,3]` would be divided by 2 twice by this algorithm, giving the result `[-0.75, 0.25, 0.75]`.

<!-- grade_id: cell-261e41e9d77d0635, points: 0, grade: false, solution: true -->
```python
def rescale_array(x):
    y = x.copy()  # Return a new rescaled array instead of mutating
    
    ### BEGIN SOLUTION
    while np.any(np.abs(y) > 1):
        y /= 2
    
    ### END SOLUTION
    
    
    return y
```

<!-- grade_id: cell-408327640e4f8bb2, points: 0, grade: true, solution: false -->
```python
y1 = rescale_array(np.linspace(0,10,11))
print(y1)  # Should range from 0 to 10/16 = 5/8 = 0.625.

y2 = rescale_array(np.array([-7.,-1.,1.,3.]))
print(y2)  # Should range from -7/8 = -0.875 to 3/8 = 0.375.

y3 = rescale_array(np.linspace(-1,1,20))
print(y3)  # Should range from -1 to 1 - no rescaling at all!
```

Output:
```
[0.     0.0625 0.125  0.1875 0.25   0.3125 0.375  0.4375 0.5    0.5625
 0.625 ]
[-0.875 -0.125  0.125  0.375]
[-1.         -0.89473684 -0.78947368 -0.68421053 -0.57894737 -0.47368421
 -0.36842105 -0.26315789 -0.15789474 -0.05263158  0.05263158  0.15789474
  0.26315789  0.36842105  0.47368421  0.57894737  0.68421053  0.78947368
  0.89473684  1.        ]
```

## T11.2 - A better square root

All the way back in tutorial 2, we introduced an algorithm for calculating the square root $y = \sqrt{x}$, which had the following steps:

1. Start with an initial guess for y.
2. Calculate y*y, and stop if it's "close enough" to x.
3. Update y to the average of y and x/y: y --> (y+x/y)/2.
4. Loop back to step 2.

At that point, we couldn't really do the last step: we just repeated the steps by hand until the result looked good (or we got sick of it.)  But now we can make Python control the loop for us!

### Part A

__Implement the function square_root()__ in the cell below: I've included the arguments and a brief docstring to get your started.  The "_error tolerance_" `tol` defines the __stopping condition__ of our algorithm, i.e. it defines what "close enough" in step 2 means:

$$
|y^2 - x| <= (\rm{tol})
$$

<!-- grade_id: cell-053729857a578d44, points: 0, grade: false, solution: true -->
```python
def square_root(x, y0, tol=1e-6):
    """
    Computes the square root of x using Heron's method.
    Returns: the square root of x.
    
    Arguments:
    =====
    x -- number to take the square root of.
    y0 -- initial guess for the square root of x.
    
    Keyword arguments:
    =====
    tol -- absolute error tolerance for the algorithm to stop (default: 1e-6.)
    """
    
    y = y0

    ### BEGIN SOLUTION
    while abs(y*y-x) > tol:
        y = 0.5*(y + x/y)

    ### END SOLUTION
    

    return y
```

```python
# Simple human test, before the automated ones: should print 1.414...
print(square_root(2,2))
```

Output:
```
1.4142135623746899
```

If the cell above doesn't print anything, and you see `In [*]:` to the left of it, then you have probably caused an infinite loop!  You can't run any other cells now, since the kernel will keep trying to finish this one forever.

__Don't panic!__  All you have to do is interrupt the kernel using `i`,`i`, restart it with `0`,`0` (both in _command mode_, so click outside the cells!), or use the "Kernel" menu above.  Then go back to your function and try to figure out where the infinite loop is coming from...the most likely culprit is a condition on your `while` loop that is always satisfied.

<!-- grade_id: cell-709404c8e5d5ef5c, points: 0, grade: true, solution: false -->
```python
import numpy.testing as npt
import math

npt.assert_allclose(square_root(4,4), 2)
npt.assert_allclose(square_root(47,4), math.sqrt(47))
npt.assert_allclose(square_root(3,1000), math.sqrt(3))
npt.assert_allclose(square_root(57.34,139), math.sqrt(57.34))
```

### Part B _(optional challenge)_

In addition to `tol`, we might like to have an absolute cap on the number of iterations: this will safeguard against infinite (and finite but really long) loops.  Implement the new and improved version `square_root_v2` below, which will also stop if the number of loop iterations reaches `max_iter`.  If `max_iter` is reached, your function should print the message:
```python
"Warning: max_iter reached!  Result may be inaccurate."
```
and then __use the `break` keyword to terminate the loop__ and return the current value of `y`.

(Hint: don't forget to update `num_iter` in your loop!)

<!-- grade_id: cell-bb87341f5e6835b6, points: 0, grade: false, solution: true -->
```python
def square_root_v2(x, y0, tol=1e-6, max_iter=50):
    """
    Computes the square root of x using Heron's method.
    Returns: the square root of x.
    
    Arguments:
    =====
    x -- number to take the square root of.
    y0 -- initial guess for the square root of x.
    
    Keyword arguments:
    =====
    tol -- absolute error tolerance for the algorithm to stop (default: 1e-6.)
    max_iter -- maximum number of iterations to run (default: 50)
    """

    y = y0
    num_iter = 0
    
    ### BEGIN SOLUTION
    while abs(y*y-x) > tol:
        if num_iter >= max_iter:
            print("Warning: max_iter reached!  Result may be inaccurate.")
            break

        y = 0.5*(y + x/y)
        num_iter += 1       
    ### END SOLUTION
    

    return y

```

<!-- grade_id: cell-595069b21ae40839, points: 0, grade: true, solution: false -->
```python
import numpy.testing as npt
import math

print(square_root_v2(3, 1000))
npt.assert_allclose(square_root_v2(4,4), 2)
npt.assert_allclose(square_root_v2(47,4), math.sqrt(47))
npt.assert_allclose(square_root_v2(3,1000, max_iter=100), math.sqrt(3))
npt.assert_allclose(square_root_v2(57.34,139, tol=1e-10), math.sqrt(57.34))

# Should see a warning message from this:
npt.assert_allclose(square_root_v2(47, 3, max_iter=2), 7.1, atol=0.1)
```

Output:
```
1.7320508075705185
Warning: max_iter reached!  Result may be inaccurate.
```

### Bonus Part C

As a little __bonus__, let's see how the performance of our pure Python function compares to the one from the math library.  The `%timeit` magic command will tell Jupyter to re-run a particular function over and over and time how long it takes.  (We can also use `%time`, which will just run it once; a single run will be less reliable for estimating the average run time, though.)

(If you didn't finish part B above, you can run the test on `square_root()` instead of `square_root_v2()`.)

```python
%timeit square_root_v2(455, 40)
%timeit np.sqrt(455)
```

Output:
```
1.47 µs ± 437 ns per loop (mean ± std. dev. of 7 runs, 1000000 loops each)
856 ns ± 284 ns per loop (mean ± std. dev. of 7 runs, 1000000 loops each)
```